# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
import getpass
myname = getpass.getuser().split('-')[1]
myname
MYTABLE=myname + '_data'
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
host = 'slucor2022-instance-1.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2022'
password = 'SLUcor2022'
database = 'hds5210'
# database = ''
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'pboal')
assert(conn.connect())
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
pd.read_sql_query("SHOW TABLES", conn)
```

Output:
```
Tables_in_hds5210
0                   data
1          ecdmdphd_data
2              ed_visits
3      hospital_services
4              hospitals
5     lillyeversman_data
6          paulboal_data
7           payment_data
8               per_data
9             population
10      sravani9544_data
11       svelicheti_data
12   system_affiliations
13  total_per_population
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
tables = {}

for n,t in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(t['Tables_in_hds5210'])
    print(name)
    tables[name] = pd.read_sql("SELECT * FROM "+name,conn)
```

Output:
```
data
ecdmdphd_data
ed_visits
hospital_services
hospitals
lillyeversman_data
paulboal_data
payment_data
per_data
population
sravani9544_data
svelicheti_data
system_affiliations
total_per_population
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
tables['population']
```

Output:
```
index    Zip  Population
0         0  90011      110750
1         1  90650      104524
2         2  91331      103683
3         3  90201      100512
4         4  90044       99980
...     ...    ...         ...
1645   1645  95699          13
1646   1646  96119          11
1647   1647  92304           7
1648   1648  93530           6
1649   1649  92364           5

[1650 rows x 3 columns]
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

https://opendata.arcgis.com/api/v3/datasets/ed0bc689bde246b18835c7529ba4efb4_0/downloads/data?format=csv&spatialRefId=4326

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
#import csv file using pandas pd 

import pandas as pd
import matplotlib.pyplot as plt

data = pd.read_csv('/home/jupyter-msimkins7/hds5210-2022/Electronic_Health_Record_EHR_Incentive_Program_Payments_to_Eligible_Hospitals.csv')
```

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
MYTABLE = pd.read_sql_query('SELECT * FROM data', con=conn)
```

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
MYTABLE.columns
```

Output:
```
Index(['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'],
      dtype='object')
```

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
MYTABLE.shape
```

Output:
```
(323, 22)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (323,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'])
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/ops/array_ops.py in masked_arith_op(x, y, op)
    110         if mask.any():
    111             with np.errstate(all="ignore"):
--> 112                 result[mask] = op(xrav[mask], y)
    113 
    114     result, _ = maybe_upcast_putmask(result, ~mask, np.nan)

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/ops/roperator.py in radd(left, right)
      7 
      8 def radd(left, right):
----> 9     return right + left
     10 
     11 

UFuncTypeError: ufunc 'add' did not contain a loop with signature matching types (dtype('<U32'), dtype('<U32')) -> dtype('<U32')
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Business_ZIP_Code` field and `population`'s `Zip` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
population = tables['population']
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer = pd.merge(MYTABLE,
                  population,
                  left_on = 'Business_ZIP_Code',
                  right_on = 'Zip',
                  how = 'inner')
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer
```

Output:
```
index_x           X          Y  OBJECTID  \
0          0 -124.142009  40.783559         1   
1          1 -122.086674  37.632915         2   
2        194 -122.096165  37.630272       195   
3          2 -122.295861  38.325402         3   
4          3 -120.766467  38.350920         4   
..       ...         ...        ...       ...   
312      312 -122.907916  39.105474       313   
313      313 -122.267252  37.837338       314   
314      316 -117.833859  33.782715       317   
315      318 -122.173214  37.436328       319   
316      321 -118.286401  34.067178       322   

                                         Provider_Name         NPI    CCN  \
0             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1609858950  50006   
1                             HAYWARD SISTERS HOSPITAL  1942298153  50002   
2                          KAISER FOUNDATION HOSPITALS  1801960513  50512   
3             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1235218785  50009   
4                              
```

```python
answer['NPI']
```

Output:
```
0      1609858950
1      1942298153
2      1801960513
3      1235218785
4      1447494323
          ...    
312    1952634008
313    1003961251
314    1225015985
315    1467442749
316    1508908591
Name: NPI, Length: 317, dtype: int64
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
max_pop = answer['Population'].max()
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
max_pop
```

Output:
```
104524
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
filter = answer['Population'] ==104524
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
max_pop_only = answer[filter]
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
max_pop_only
```

Output:
```
index_x          X          Y  OBJECTID                 Provider_Name  \
280      275 -118.09915  33.912613       276  CPH HOSPITAL MANAGEMENT, LLC   

            NPI    CCN Business_Street_Address Business_City Business_County  \
280  1194016923  50771     13100 Studebaker Rd       Norwalk     Los Angeles   

     ...  total_payments Last_Payment_Criteria  Recent_Disbursement_Amount  \
280  ...       2184880.1                    MU                   218488.01   

      Latitude  Longitude  Last_Program_Year Last_Payment_Year  index_y  \
280  33.912613 -118.09915               2018              2020        1   

       Zip  Population  
280  90650      104524  

[1 rows x 25 columns]
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer = []
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer.append(max_pop_only[['NPI']])
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
answer
```

Output:
```
[            NPI
 280  1194016923]
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

Output:
```
<ipython-input-212-fc652b179429> in <module>
      1 assert(type(answer) == list)
----> 2 assert(answer == ['1194016923'])

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/generic.py in __nonzero__(self)
   1328     def __nonzero__(self):
   1329         raise ValueError(
-> 1330             f"The truth value of a {type(self).__name__} is ambiguous. "
   1331             "Use a.empty, a.bool(), a.item(), a.any() or a.all()."
   1332         )

ValueError: The truth value of a DataFrame is ambiguous. Use a.empty, a.bool(), a.item(), a.any() or a.all().
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

## Now using Pandas instead of SQL

```python
import numpy as np
test = data.pivot_table(
    values = 'total_payments',
    index= ['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type'],
    #columns = 'Medicaid_EP_Hospital_Type',
    aggfunc = np.ma.mean,
    fill_value = 0)
```

```python
test.reset_index()
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  total_payments
0                90004      Acute Care Hospitals      1093021.73
1                90012      Acute Care Hospitals      3614739.93
2                90015      Acute Care Hospitals     10045484.04
3                90017      Acute Care Hospitals      3899845.21
4                90020      Children's Hospitals       427640.22
..                 ...                       ...             ...
262              96104      Acute Care Hospitals       184615.39
263              96122      Acute Care Hospitals       531907.08
264              96130      Acute Care Hospitals      1371962.88
265              96150      Acute Care Hospitals       786769.13
266              96161      Acute Care Hospitals       621315.39

[267 rows x 3 columns]
```

```python
#merge data with population data
new_test = pd.merge(test.reset_index(),
                  population,
                  left_on = 'Business_ZIP_Code',
                  right_on = 'Zip',
                  how = 'inner')
```

```python
new_test
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  total_payments  index  \
0                90004      Acute Care Hospitals      1093021.73    115   
1                90012      Acute Care Hospitals      3614739.93    433   
2                90015      Acute Care Hospitals     10045484.04    738   
3                90017      Acute Care Hospitals      3899845.21    648   
4                90020      Children's Hospitals       427640.22    390   
..                 ...                       ...             ...    ...   
256              96104      Acute Care Hospitals       184615.39   1388   
257              96122      Acute Care Hospitals       531907.08   1169   
258              96130      Acute Care Hospitals      1371962.88    810   
259              96150      Acute Care Hospitals       786769.13    595   
260              96161      Acute Care Hospitals       621315.39    826   

       Zip  Population  
0    90004       61105  
1    90012       37268  
2    90015       23900  
3   
```

```python
#calculate total payment per person based on zip population
new_test['total_per_person'] = (new_test["total_payments"])/new_test['Population']
```

```python
new_test
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  total_payments  index  \
0                90004      Acute Care Hospitals      1093021.73    115   
1                90012      Acute Care Hospitals      3614739.93    433   
2                90015      Acute Care Hospitals     10045484.04    738   
3                90017      Acute Care Hospitals      3899845.21    648   
4                90020      Children's Hospitals       427640.22    390   
..                 ...                       ...             ...    ...   
256              96104      Acute Care Hospitals       184615.39   1388   
257              96122      Acute Care Hospitals       531907.08   1169   
258              96130      Acute Care Hospitals      1371962.88    810   
259              96150      Acute Care Hospitals       786769.13    595   
260              96161      Acute Care Hospitals       621315.39    826   

       Zip  Population  total_per_person  
0    90004       61105         17.887599  
1    90012       3
```

```python
import numpy as np
answer = new_test.pivot_table(
    values = 'total_per_person',
    index= 'Medicaid_EP_Hospital_Type',
    aggfunc = np.ma.mean,
    fill_value = 0)
```

```python
answer
```

Output:
```
total_per_person
Medicaid_EP_Hospital_Type                  
Acute Care Hospitals              85.444853
Children's Hospitals             315.394944
```

```python

```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/frame.py in __getitem__(self, key)
   2904             if self.columns.nlevels > 1:
   2905                 return self._getitem_multilevel(key)
-> 2906             indexer = self.columns.get_loc(key)
   2907             if is_integer(indexer):
   2908                 indexer = [indexer]

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/indexes/base.py in get_loc(self, key, method, tolerance)
   2898                 return self._engine.get_loc(casted_key)
   2899             except KeyError as err:
-> 2900                 raise KeyError(key) from err
   2901 
   2902         if tolerance is not None:

KeyError: 'Avg_Pay_per_Capita'
```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week15_assignment.ipynb
    !git commit -a -m "Submitting the week 15 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

```python

```